### Graded Lab 4

Hello ! Welcome to Graded Lab of Module 4.

In the last assignment we had worked on correlation , hypothesis which business wanted us to test.
Its high time when we look at the missing values & outliers stuff.
**Here we will work on missing values & outliers.**

In case if you are not able to recollect the problem description and data description then mentioning it below.

Lets look at the problem statement,

*Client: ABC Retail, Incorporated, rest-of-the-world division* 

***Project name: Online retail sales analysis*** 

An online retailer, ABC, Inc., operates in nearly 100 countries worldwide, selling furniture, office supplies and technology products to customers in three segments: consumer, corporate and home office. ABC, Inc. is a US-based company, and it has two major divisions: US and rest of the world. We are working with the rest of the world division of the company. 

They have provided us with online sales transaction data from 2011 to 2014.

We are given 3 datasets:-

1. Data on each sale; 51290 records; all data in US dollars
It contains fields like
**order_id** (identifier) ,order_date ,ship_date ,ship_mode ,**customer_id**(identifier) ,product_id ,category ,sub_category ,product_name ,sales ,quantity ,discount ,profit ,shipping_cost ,order_priority ,**vendor_code** (identifier) 


2. Data on the customers; 1590 records 
It contains fields like
**customer_id** (identifier) ,customer_name ,city ,state ,country ,postal_code ,segment ,market ,region 

3. Data on vendors who supply the retailer; 65 records 
It contains fields like
vendor ,**vendor_code** (identifier) 

We need to analyze the data and need to provide answer to different questions asked by company officials.

```python
# importing libraries
import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
from scipy.stats import norm
from itertools import combinations
from sklearn.preprocessing import OrdinalEncoder, LabelEncoder

### Reading sales data
sales = pd.read_csv('sales_data_M4.csv')

### Reading customer data
cust = pd.read_csv(r'customers.csv',encoding='iso-8859-1')

### Reading vendor data
vend = pd.read_csv(r'vendors.csv')

sales.head()
```

Output:
```
row_id        order_id  order_date   ship_date       ship_mode customer_id  \
0       1  MX-2014-143658  02/10/2014  06/10/2014  Standard Class    SC-20575   
1       2  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
2       3  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
3       4  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
4       5  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   

        product_id         category sub_category  \
0  OFF-LA-10002782  Office Supplies       Labels   
1  FUR-FU-10004015        Furniture  Furnishings   
2  FUR-BO-10002352        Furniture    Bookcases   
3  OFF-BI-10004428  Office Supplies      Binders   
4  OFF-AR-10004594  Office Supplies          Art   

                         product_name   sales  quantity  discount  profit  \
0  Hon File Folder Labels, Adjustable   13.08         3       0.0    4.56   
1                Tenex Clock, Durable  252.16    
```

In order to solve the next questions , we need to combine all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. So here we have written a data processing function.
There are 2 tasks which are to be performed.
1. Merge/ Join all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. (Understand which should be the joining key , type of join , refer .merge() function of pandas)
2. Convert 'order_date' into a datetime column.
**Return output as a dataframe**

<!-- grade_id: cell-760043ec960a09cb, points: 0, grade: false, solution: true -->
```python
#### data_merging & order_date processing , data1 will be sales , data2 will be customer dataset & data3 will be vendor dataset.

def data_process(data1,data2,data3):
    # your code here
    data = pd.merge(data1,data2,on='customer_id',how='left')
    data = pd.merge(data,data3,on='vendor_code',how='left')
    data['order_date'] = pd.to_datetime(data['order_date'])
    return data
```

```python
sales= data_process(data1=sales.copy(),data2=cust.copy(),data3=vend.copy())
```

```python
sales.sample(100).head()
```

Output:
```
row_id         order_id order_date   ship_date       ship_mode  \
16683   16684  ES-2012-5959969 2012-10-30  04/11/2012  Standard Class   
29127   29128    ID-2012-22550 2012-11-13  17/11/2012  Standard Class   
20517   20518    ID-2014-53679 2014-10-29  05/11/2014  Standard Class   
48261   48262     MZ-2014-4670 2014-06-24  29/06/2014  Standard Class   
47415   47416      BO-2014-170 2014-06-25  01/07/2014  Standard Class   

      customer_id        product_id         category sub_category  \
16683    AJ-10780   OFF-AP-10004869  Office Supplies   Appliances   
29127    BD-11620   OFF-ST-10003931  Office Supplies      Storage   
20517    PP-18955   TEC-PH-10004618       Technology       Phones   
48261      AZ-750  OFF-IBI-10004855  Office Supplies      Binders   
47415    TC-11145  FUR-TEN-10003349        Furniture  Furnishings   

                              product_name  ...  vendor_code    customer_name  \
16683            KitchenAid Toaster, White  ...       VE_018   Anthony J
```

<!-- grade_id: cell-24e204e9afc2c313, points: 0, grade: true, solution: false -->
```python
assert sales['order_date'].dtypes=='<M8[ns]' ,'Make sure if you have converted order_date into a datetime format correctly or not.'
assert sales.shape== (51290,26) ,'Checking size and shape of dataframe after merging is a very important check.'
```

Note that in this dataset we are purposefully introducing some of the missing values.

### Q.1 Return column names & missing values percentage for the columns which has missing value% >50. Output should be a dictionary. 

For eg:- { column_name : value%} i.e {'A':54.50,'B':59.75}, make sure to round off percentages to 2 decimals.

<!-- grade_id: cell-eaa4ef1cb3c93902, points: 0, grade: false, solution: true -->
```python
def missing_value_col(data):
    # your code here
    missing_data = data.isnull().sum()/len(data)*100
    missing_data = missing_data.to_dict()
    missing_cols = {key:missing_data[key] for key in missing_data.keys() if missing_data[key]>50}
    return missing_cols
```

```python
missing_value_col(data=sales)
```

Output:
```
{'postal_code': 100.0}
```

```python
list(missing_value_col(data=sales).keys())
```

Output:
```
['postal_code']
```

<!-- grade_id: cell-6d0d4c99764805c4, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

#### **Senior data scientists after consulting with the business have decided that such columns needs to be dropped before going further.**

```python
high_missing_col=list(missing_value_col(data=sales).keys()) #Enter the columns with high % of missing values
### We will drop columns with such high missing values
sales.drop(high_missing_col,axis=1,inplace=True)
```

```python
sales.columns
```

Output:
```
Index(['row_id', 'order_id', 'order_date', 'ship_date', 'ship_mode',
       'customer_id', 'product_id', 'category', 'sub_category', 'product_name',
       'sales', 'quantity', 'discount', 'profit', 'shipping_cost',
       'order_priority', 'vendor_code', 'customer_name', 'city', 'state',
       'country', 'segment', 'market', 'region', 'vendor'],
      dtype='object')
```

### Q.2 Fill the missing data in 'category' column with appropriate method.  
(Hint:- Remember variable 'category' & 'sub-category' are related & hierachy goes like this category -> sub-category)

Approach:-

1. Create a key-pair dictionary where key will be product category & pair values will be sub-categories.

2. Wherever category value is missing check for corresponding sub-category value. 

3. From sub-category value trace back its original product category value by looking in dictionary, get a product value.

4. Fill the blank product category value by the same.

```python
sales.sample(100).head(2)
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode  \
9578     9579  MX-2011-155257 2011-07-13  15/07/2011     First Class   
21122   21123   IN-2012-74686 2012-12-17  21/12/2012  Standard Class   

      customer_id       product_id         category sub_category  \
9578     JH-16180  OFF-ST-10001172              NaN      Storage   
21122    JH-15910  OFF-LA-10000471  Office Supplies       Labels   

                                product_name  ...  order_priority  \
9578              Fellowes Trays, Wire Frame  ...            High   
21122  Smead Shipping Labels, 5000 Label Set  ...            High   

       vendor_code    customer_name         city             state    country  \
9578        VE_006     Justin Hirsh  Mexico City  Distrito Federal     Mexico   
21122       VE_030  Jonathan Howell    Matagalpa         Matagalpa  Nicaragua   

        segment market         region       vendor  
9578   Consumer  LATAM    LATAM-North  Venusystems  
21122  Consumer  LATAM  LAT
```

<!-- grade_id: cell-21490e58b10fdc96, points: 0, grade: false, solution: true -->
```python
### create a dictionary where keys will be categories & sub-categories as values. store in inside Category_dict.

def cat_dict(data,main_col,sub_col):    
    # your code here
    categories = data[[main_col,sub_col]].drop_duplicates()
    Category_dict = categories.groupby(main_col)[sub_col].apply(list).to_dict()
    return Category_dict

def category_subcategory_map_dataframe(df, Category_dict, main_col, sub_col):
    def category_subcategory_map(row):
        if pd.isna(row[main_col]):  # If the main column is null
            if not pd.isna(row[sub_col]):  # If the sub column is not null
                for k, v in Category_dict.items():
                    if row[sub_col] in Category_dict[k]:
                        return k  # k is the correct sub-category for the category value v
        else:
            return row[main_col]  # If the main column is not null, return its existing value

    df[main_col] = df.apply(category_subcategory_map, axis=1)
    return df[main_col]
```

```python
cat_dict(data=sales,main_col='category',sub_col='sub_category')
```

Output:
```
{'Furniture': ['Furnishings', 'Bookcases', 'Chairs', 'Tables'],
 'Office Supplies': ['Labels',
  'Binders',
  'Art',
  'Envelopes',
  'Supplies',
  'Fasteners',
  'Storage',
  'Appliances',
  'Paper'],
 'Technology': ['Machines', 'Accessories', 'Copiers', 'Phones']}
```

```python
category_dict=cat_dict(data=sales,main_col='category',sub_col='sub_category')
category_subcategory_map_dataframe(df=sales, Category_dict=category_dict, main_col='category', sub_col='sub_category')
```

Output:
```
0        Office Supplies
1              Furniture
2              Furniture
3        Office Supplies
4        Office Supplies
              ...       
51285    Office Supplies
51286         Technology
51287    Office Supplies
51288    Office Supplies
51289    Office Supplies
Name: category, Length: 51290, dtype: object
```

```python
Category_dict=cat_dict(data=sales,main_col='category',sub_col='sub_category')
sales['category']=category_subcategory_map_dataframe(df=sales, Category_dict=Category_dict, main_col='category', sub_col='sub_category')
```

<!-- grade_id: cell-b7f7ba970eee54ec, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.3 ) Fill the missing values in 'sales' column by average value of that column. Return output series.

<!-- grade_id: cell-7ef02c9d67c20af2, points: 0, grade: false, solution: true -->
```python
def fill_missing(data,col):
    # your code here
    data[col].fillna(data[col].mean(),inplace=True)
    return data[col]
```

```python
sales.isna().sum()
```

Output:
```
row_id               0
order_id             0
order_date           0
ship_date            0
ship_mode         1570
customer_id          0
product_id           0
category             0
sub_category         0
product_name      1297
sales             6204
quantity             0
discount             0
profit               0
shipping_cost        0
order_priority    1046
vendor_code          0
customer_name        0
city                 0
state                0
country              0
segment              0
market               0
region               0
vendor               0
dtype: int64
```

```python
fill_missing(data=sales,col='sales').isna().sum()
```

Output:
```
0
```

```python
sales['sales']=fill_missing(data=sales,col='sales')
```

<!-- grade_id: cell-519682f18719b5ec, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
assert sales['sales'].isna().sum()==0 ,'Missing values found for sales.'
```

### Q.4 ) Discretize the 'sales' column wrto following groups. Each group should contain number of orders as values.
Groups :- [0,12),[12,37),[37,113),[113,186),[186,1200).

**Note :-[a,b) indicates that group includes a but excudes b.**

**Return output series with group name as index & number of orders as values.**

For eg:-

[186, 1200)             123

[37, 113)               123

[12, 37)                123

[113, 186)              123

[0, 12)                 123

Name: sales, dtype: int64

```python
sales.sample(100).head()
```

Output:
```
row_id         order_id order_date   ship_date       ship_mode  \
32176   32177   CA-2014-111689 2014-01-12  03/12/2014    Second Class   
32286   32287   CA-2012-153549 2012-03-29  31/03/2012    Second Class   
31647   31648   CA-2013-129714 2013-02-09  04/09/2013     First Class   
15641   15642  ES-2011-2535754 2011-08-03  13/03/2011  Standard Class   
50117   50118     AG-2011-8180 2011-02-25  25/02/2011        Same Day   

      customer_id        product_id         category sub_category  \
32176    HP-14815   FUR-CH-10004287        Furniture       Chairs   
32286    SL-20155   FUR-CH-10004086        Furniture       Chairs   
31647    AB-10060   OFF-BI-10002160  Office Supplies      Binders   
15641    LW-17215   OFF-FA-10000670  Office Supplies    Fasteners   
50117     RO-9780  TEC-EPS-10004558       Technology     Machines   

                                            product_name  ...  order_priority  \
32176                           SAFCO Arco Folding Chair  ...        Cri
```

<!-- grade_id: cell-011ab56604d6dc6d, points: 0, grade: false, solution: true -->
```python
def manual_bin(data,col):
    # your code here
    bin_edges = [0,12,37,113,186,1200]
    data[col+'_bin'] = pd.cut(data[col],\
                             bins = bin_edges,\
                              right = False,
                             include_lowest = True)
    return data[col+'_bin'].value_counts()
```

```python
manual_bin(data=sales,col='sales')
```

Output:
```
[186, 1200)    18550
[37, 113)      12761
[12, 37)        9325
[113, 186)      5187
[0, 12)         3708
Name: sales_bin, dtype: int64
```

<!-- grade_id: cell-c70daab6a00fb4e5, points: 5, grade: true, solution: false -->
```python
assert type(manual_bin(data=sales,col='sales'))==pd.Series, "Please provide output in series format."
```

<!-- grade_id: cell-5553c36ce34800ac, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
assert manual_bin(data=sales,col='sales').values[0]==18550, 'Make sure you have returned correct series values.'
```

### Q.5 )  Sometimes, it is difficult to comprehend the quantity sold ('quantity') value as a numeric value. So, the business team wants to observe the sales as High, Medium and Low categories.  

### Write a code to discretize the sales data and  return an output series containing High, Medium and Low categories as index & the respective number of records as values.

Hint : Here the bounds for High , Medium , Low categories are not given hence we can consider the equal split of data for defining the bounds , think of using pd.cut() function.

```python
# turning the numeric data in discretized data
data = pd.cut(sales['quantity'], bins = 3,include_lowest=True)#,labels = ['Low', 'Medium', 'High'])
data.value_counts()
```

Output:
```
(0.986, 5.333]    42660
(5.333, 9.667]     7753
(9.667, 14.0]       877
Name: quantity, dtype: int64
```

```python
data = pd.cut(sales['quantity'], bins = 3,include_lowest=True,labels = ['Low', 'Medium', 'High'])
data.value_counts()
```

Output:
```
Low       42660
Medium     7753
High        877
Name: quantity, dtype: int64
```

<!-- grade_id: cell-8e4399b1bccc98dd, points: 0, grade: false, solution: true -->
```python
def discrete_bin(data,col):
    # your code here
    data_bin = pd.cut(data[col], bins = 3,include_lowest=True,labels = ['Low', 'Medium', 'High'])
    return data_bin.value_counts()
```

<!-- grade_id: cell-8f418bc782917abb, points: 5, grade: true, solution: false -->
```python
assert type(discrete_bin(data=sales,col='quantity'))==pd.Series, "Please provide output in series format."
```

<!-- grade_id: cell-4a156f92f5438b15, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.6. In order to perform further analysis, we intend to utilize the 'ship_mode' variable, which has categorical values ('First Class'(0) > 'Second Class'(1) > 'Standard Class'(2) > 'Same Day'(3)). To facilitate this analysis, we need to convert 'ship_mode' into a numerical feature. Please encode 'ship_mode' into a numerical format.
### Return output series.

Hint : Think of using the OrdinalEncoder() function from scikit-learn.

```python
sales[['ship_mode','order_priority']].head()
```

Output:
```
ship_mode order_priority
0  Standard Class         Medium
1  Standard Class         Medium
2  Standard Class         Medium
3  Standard Class         Medium
4  Standard Class         Medium
```

```python
sales['order_priority'].unique()
```

Output:
```
array(['Medium', 'Low', 'High', nan, 'Critical'], dtype=object)
```

```python
sales[['ship_mode']]
```

Output:
```
ship_mode
0      Standard Class
1      Standard Class
2      Standard Class
3      Standard Class
4      Standard Class
...               ...
51285     First Class
51286             NaN
51287     First Class
51288     First Class
51289  Standard Class

[51290 rows x 1 columns]
```

```python
priority = ['First Class','Second Class','Standard Class','Same Day','nan_val']
nan_val=5
from sklearn.preprocessing import OrdinalEncoder
import numpy as np
encoder = OrdinalEncoder(categories=[priority])
data = sales.copy()
col = 'ship_mode'
data['ship_mode'].fillna('nan_val',inplace=True)
enc_shipmode = encoder.fit_transform(data[['ship_mode']])
enc_shipmode = np.where(enc_shipmode==4,nan_val,enc_shipmode)
enc_shipmode[-5:]
```

Output:
```
array([[0.],
       [5.],
       [0.],
       [0.],
       [2.]])
```

<!-- grade_id: cell-d684beb2bc73d5a6, points: 0, grade: false, solution: true -->
```python
## Select Appropriate priority
priority = ['First Class','Second Class','Standard Class','Same Day','nan_val'] # Write your code in place of None 
def transforming_ordered_var(data,col,priority,nan_val):
    #data: dataset to be used
    #col : column name which is to be encoded.
    #priority: list of categories sorted from highest order of importance to lowest order importance.
    #nan_val : it will be a value assigned in place of nan's. 
    #
    # your code here
    from sklearn.preprocessing import OrdinalEncoder
    import numpy as np
    import pandas as pd
    encoder = OrdinalEncoder(categories=[priority])
#     data = sales.copy()
#     col = 'ship_mode'
    data['ship_mode'].fillna('nan_val',inplace=True)
    enc_shipmode = encoder.fit_transform(data[['ship_mode']])
    enc_shipmode = enc_shipmode.reshape(-1)
    enc_shipmode = np.where(enc_shipmode==4,nan_val,enc_shipmode)
    enc_shipmode = pd.Series(enc_shipmode)
    return enc_shipmode
```

```python
sales['shipping_mode_encoded']=transforming_ordered_var(data=sales,col='ship_mode',priority=priority,nan_val=5)
```

<!-- grade_id: cell-f1ba60c8543b3583, points: 20, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
assert int(transforming_ordered_var(data=sales,col='ship_mode',priority=priority,nan_val=5).value_counts().loc[lambda x:x.index==5])==1570
```

### Q.7) Outliers can be defined in different ways. For each of the following definitions of outliers, what percent of the values in the  'sales' are considered outliers? Give the answer as a percent rounded upto two decimal places.


Note -
1. 'IQR_detection' : Less than (Q1 - 1.5*IQR) and greater than (Q3 + 1.5*IQR) are considered outliers

2. 'percentile_detection': Below the 3rd percentile and above the 97th percentile are considered outliers

3. 'mean_SD' : Less than (mean - 3*SD) and greater than (mean + 3*SD) are considered outliers 

```python
data = sales.copy()
col = 'sales'
Q1 = data[col].quantile(0.25)
Q3 = data[col].quantile(0.75)
IQR = Q3 - Q1
outliers = data[(data[col] < (Q1 - 1.5 * IQR)) | (data[col] > (Q3 + 1.5 * IQR))]
outliers_perc = round(len(outliers)/len(sales),2)
```

```python
data = sales.copy()
col = 'sales'
outliers = data[(data[col] < data[col].quantile(0.03)) | (data[col] > data[col].quantile(0.97))]
outliers_perc = round(len(outliers)*100/len(sales),2)
```

```python
data = sales.copy()
col = 'sales'
mean = data[col].mean()
SD = data[col].std()
outliers = data[(data[col] < (mean - 3*SD)) | (data[col] > (mean + 3*SD))]
outliers_perc = round(len(outliers)*100/len(sales),2)
```

<!-- grade_id: cell-c1830537e6d495e9, points: 0, grade: false, solution: true -->
```python
def detect_outliers(data,col,method):
    if method=='IQR_detection':
        # your code here
        Q1 = data[col].quantile(0.25)
        Q3 = data[col].quantile(0.75)
        IQR = Q3 - Q1
        outliers = data[(data[col] < (Q1 - 1.5 * IQR)) | (data[col] > (Q3 + 1.5 * IQR))]
        outliers_perc = round(len(outliers)*100/len(sales),2)
    elif method=='percentile_detection':
        # your code here
        outliers = data[(data[col] < data[col].quantile(0.03)) | (data[col] > data[col].quantile(0.97))]
        outliers_perc = round(len(outliers)*100/len(sales),2)
    elif method=='mean_SD':
        # your code here
        mean = data[col].mean()
        SD = data[col].std()
        outliers = data[(data[col] < (mean - 3*SD)) | (data[col] > (mean + 3*SD))]
        outliers_perc = round(len(outliers)*100/len(sales),2)
    return outliers_perc
```

<!-- grade_id: cell-618cb2f39f3bdbb8, points: 5, grade: true, solution: false -->
```python
assert type(detect_outliers(data=sales,col='sales',method='IQR_detection'))==np.float, 'Make sure that output is returned as float'
assert detect_outliers(data=sales,col='sales',method='IQR_detection')==10.06, 'Make sure that output is rounded upto 2 decimals'
```

<!-- grade_id: cell-bfd2679b1e2a3929, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.8) Detect upper end outliers in profit column by 'IQR_detection'. Return number of upper end outliers.

```python
data = sales.copy()
col = 'profit'
Q1 = data[col].quantile(0.25)
Q3 = data[col].quantile(0.75)
IQR = Q3 - Q1
upper_outliers = data[data[col] > (Q3 + 1.5 * IQR)]
upper_outliers.shape[0]
```

Output:
```
6229
```

<!-- grade_id: cell-96b2becd9fec86dd, points: 0, grade: false, solution: true -->
```python
def detect_upperend_outliers(data,col):
    # your code here
    Q1 = data[col].quantile(0.25)
    Q3 = data[col].quantile(0.75)
    IQR = Q3 - Q1
    outliers = data[data[col] > (Q3 + 1.5 * IQR)]
    outliers.shape
    return outliers.shape[0] ##data frame named 'outliers' which contains only outliers.
```

<!-- grade_id: cell-9623402867c6a9ef, points: 5, grade: true, solution: false -->
```python
assert type(detect_upperend_outliers(data=sales,col='profit'))==int , 'No of outliers should be an integer , make sure that you are returning output in integer format.'
```

<!-- grade_id: cell-5e43e81906b9a713, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.9) Cap & Floor outliers at 97% at 3%. Return the processed column as output.

```python
data = sales.copy()
col = 'profit'
lower_limit = data[col].quantile(0.03)
upper_limit = data[col].quantile(0.97)
capped_col = data[col]
print(lower_limit,upper_limit)
capped_col = np.where(capped_col<=lower_limit,lower_limit,capped_col)
capped_col = np.where(capped_col>=upper_limit,upper_limit,capped_col)
print(np.quantile(capped_col,0.03),np.quantile(capped_col,0.97))
```

Output:
```
-144.63594 301.97196
-144.6319602 301.9693068
```

<!-- grade_id: cell-f5527eddcb6c9757, points: 0, grade: false, solution: true -->
```python
def cap_and_floor_column(column, cap_percentile, floor_percentile):
    # your code here
    lower_limit = data[col].quantile(floor_percentile/100)
    upper_limit = data[col].quantile(cap_percentile/100)
    capped_col = data[col]
    capped_col = np.where(capped_col<=lower_limit,lower_limit,capped_col)
    capped_col = np.where(capped_col>=upper_limit,upper_limit,capped_col)
    return capped_col ## Capped ,Floored series
```

```python
sales['profit_outlier_capped'] = cap_and_floor_column(column=sales['profit'], cap_percentile=97, floor_percentile=3)
```

<!-- grade_id: cell-e8e8dbefcd633180, points: 5, grade: true, solution: false -->
```python
assert round(np.quantile(sales['profit_outlier_capped'],0.97),4)==301.9693, 'Make sure that you are flooring and capping with correct percentiles & returning the correct output.'
assert round(np.quantile(sales['profit_outlier_capped'],0.03),4)==-144.632, 'Make sure that you are flooring and capping with correct percentiles & returning the correct output.'
```

<!-- grade_id: cell-d317f90221447667, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.10) Business team is interested in knowing how many vendors have profit_outlier_capped <10000 & vendors with profit_outlier_capped >10000. Return the dataframe with 3 columns vendor  , profit_outlier_capped , profit_bins & sort it with 'profit_outlier_capped' in descending order.

Instructions:-
1. Bins should be in (a,b] format. (a,b] indicates that bin excludes a but incudes b.
2. Make sure that for **1st interval only** it should be **left inclusive** i.e [a,b].
3. Segment the ''profit_outlier_capped'' variable in 2 bins. (binendpoints should be min value of profit<10000,10000,max value of profit>10000).

```python
data = sales.copy()
col = 'profit_outlier_capped'
maxx = data['profit_outlier_capped'].max()
minn = data['profit_outlier_capped'].min()
print(minn,maxx)
bin_edges = [minn,100,maxx]
data['profit_bins'] = pd.cut(data[col],\
                         bins = bin_edges,\
                          right = True,
                         include_lowest = True)

data1 = data[['vendor','profit_outlier_capped','profit_bins']]
data1 = data1.sort_values(by=['profit_outlier_capped'],axis=0,ascending=False)
data1.head()
```

Output:
```
-144.63594 301.97196
               vendor  profit_outlier_capped       profit_bins
24861      Gnomestone              301.97196  (100.0, 301.972]
20520     Venusystems              301.97196  (100.0, 301.972]
20686  Sail Furniture              301.97196  (100.0, 301.972]
7649             Nero              301.97196  (100.0, 301.972]
7653   Energy Systems              301.97196  (100.0, 301.972]
```

<!-- grade_id: cell-77f6b02474e37c07, points: 0, grade: false, solution: true -->
```python
def profit_vendor(data,col):
    # your code here
    maxx = data['profit_outlier_capped'].max()
    minn = data['profit_outlier_capped'].min()
    print(minn,maxx)
    bin_edges = [minn,100,maxx]
    data['profit_bins'] = pd.cut(data[col],\
                             bins = bin_edges,\
                              right = True,
                             include_lowest = True)

    data = data[['vendor','profit_outlier_capped','profit_bins']]
    data = data.sort_values(by=['profit_outlier_capped'],axis=0,ascending=False)
    return data
```

<!-- grade_id: cell-e3d429861a4da377, points: 5, grade: true, solution: false -->
```python
assert type(profit_vendor(data=sales,col='profit_outlier_capped'))==pd.DataFrame ,'Make sure that your output is a dataframe '
assert str(profit_vendor(data=sales,col='profit_outlier_capped')['profit_bins'].unique()[0])=='(10000.0, 76469.853]' , 'Make sure that you are selecting bin cut-points(a & b) properly'
assert str(profit_vendor(data=sales,col='profit_outlier_capped')['profit_bins'].unique()[1])=='(2732.9049999999997, 10000.0]', 'Make sure that you are selecting bin cut-points(a & b) properly'
```

Output:
```
-144.63594 301.97196
-144.63594 301.97196
AssertionError                            Traceback (most recent call last)
<ipython-input-201-01304fc34b2a> in <module>
      1 assert type(profit_vendor(data=sales,col='profit_outlier_capped'))==pd.DataFrame ,'Make sure that your output is a dataframe '
----> 2 assert str(profit_vendor(data=sales,col='profit_outlier_capped')['profit_bins'].unique()[0])=='(10000.0, 76469.853]' , 'Make sure that you are selecting bin cut-points(a & b) properly'
      3 assert str(profit_vendor(data=sales,col='profit_outlier_capped')['profit_bins'].unique()[1])=='(2732.9049999999997, 10000.0]', 'Make sure that you are selecting bin cut-points(a & b) properly'

AssertionError: Make sure that you are selecting bin cut-points(a & b) properly
```

<!-- grade_id: cell-05010ee8edfc15a7, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```